Python Einführungskurs für das Physikalische Anfängerpraktikum der Universität Heidelberg | [Startseite](index.ipynb)

---

# 202 - Plots mit Matplotlib

- [Einfaches Plotten](#Einfaches-Plotten)
- [Plots gestalten](#Plots-gestalten)
- [Mehrere Plots in einer Abbildung](#Mehrere-Plots-in-einer-Abbildung)
- [Plots speichern](#Plots-speichern)
- [Aufgabe 1 - Temperaturen in Heidelberg visualisieren](#Aufgabe-1---Temperaturen-in-Heidelberg-visualisieren)

Wir können mit Numpy nun numerisch Daten auswerten, und um diese zu visualisieren können wir mit dem **Matplotlib** Modul Plots von wissenschaftlicher Qualität erstellen.

Per Konvention importieren wir das PyPlot Submodul von Matplotlib unter dem Namen `plt`:

```python
# Zeige Plots direkt im Jupyter Notebook an
%matplotlib inline
# Anschließend können wir das Modul importieren
import matplotlib.pyplot as plt # Die Abkürzung `plt` ist Konvention
# Numpy brauchen wir immer.
import numpy as np
```

> **Hinweis:** Matplotlib kann mit verschiedenen _Backends_ arbeiten um Plots zu speichern oder auch interaktiv anzuzeigen. Der Aufruf `%matplotlib inline` konfiguriert das Modul für die **statische Anzeige** von Plots im Jupyter Notebook und muss **vor** dem Import von PyPlot ausgeführt werden. Dabei können auch Anzeigeparameter gesetzt werden:
>
> ```python
> %matplotlib inline
> import matplotlib
> # Setzte Anzeigeparameter, z.B.:
> matplotlib.rcParams['figure.figsize'] = (10.0, 8.0)
> import matplotlib.pyplot as plt
> ```
>
> Folgender Aufruf konfiguriert das Jupyter Notebook alternativ für die Verwendung des **interaktiven Backends**:
>
> ```python
> %matplotlib notebook
> ```
>
> Versucht's mal!

## Einfaches Plotten

PyPlot besitzt ein einfaches prozedurales Interface, mit denen wir mit nur einem Befehl Daten plotten können. Da dieses jedoch schnell an seine Grenzen kommt, wollen hier jedoch gleich das später unumgängliche, "schönere", objekt-orientierte Interface benutzen. Wer nun ein besonders schwieriges Kapitel fürchtet, der sei beruhigt, die zusätzliche Komplexität des objekt-orientierten Interface beschränkt sich für uns darauf, dass wir vor jedem Plot zwei weitere Befehle absetzen müssen:

- Um einen Plot zu erstellen brauchen wir zunächst zwei Objekte, die wir erstellen müssen. Eines repräsentiert die Abbildung, in das unser Plot gezeichnet werden wird, das andere repräsentiert einen Axenplot:

```python
abb1 = plt.figure()
axes1 = abb1.add_subplot(1,1,1) # Für's erste 
```

Unser plot-Objekt besitzt eine Vielzahl an Methoden, mit denen wir unseren Plot füllen und anpassen können
- **`plot`** füllt unseren Plot mit einer Kurve. Die Methode nimmt x- und y-Daten sowie eine Vielzahl von optionalen Argumenten zur Konfiguration an. Fehlen die x-Daten, werden die Indizes der y-Daten verwendet:

```python
axes1.plot(np.arange(100)**2)

abb1 # weist Jupyter an, das Bild abb1 im Notebook anzuzeigen
```

Ein weitere Kurve kann einfach durch einen weiteren Aufruf von **`plot`** zu unserem bestehenden Plot hinzugefügt werden.

```python
axes1.plot(100 * np.arange(100))
abb1
```

Entfernt das '`#`'-Zeichen und schaut euch die Argumente in der Dokumentation an, mit denen wir den Plot konfigurieren können:

```python
#axes1.plot?
```

Matplotlib bietet unzählige Möglichkeiten, das Aussehen von Plots anzupassen! Ihr könnt oben einige ausprobieren, wie bspw. `color`, `linestyle` und `linewidth`.

- **`scatter`** plottet Datenpunkte statt einer Linie:

```python
abb2 = plt.figure()
axes2 = abb2.add_subplot(1,1,1)

axes2.scatter(np.arange(10), np.arange(10)**2)
```

```python
#axes2.scatter?
```

- **`errorbar`** nimmt zusätzlich Fehler in x- und y-Richtung an und zeichnet Fehlerbalken:

```python
abb3 = plt.figure()
axes3 = abb3.add_subplot(1,1,1)

axes3.errorbar(np.arange(10), np.arange(10)**2, yerr=np.arange(10))
```

```python
#axes3.errorbar?
```

- **`hist`** plottet schnell ein Histogramm:

```python
abb4 = plt.figure()
axes4 = abb4.add_subplot(1,1,1)

_ = axes4.hist(np.random.normal(size=100), bins=10)
```

```python
#axes4.hist?
```

> **Hinweis:** Alternativ könnt ihr Histogramme mit `numpy.histogram` und `matplotlib.axes.Axes.bar` plotten. Diese Methode bietet etwas mehr Kontrolle über die Berechnung der Histogrammdaten.

- **`imshow`** visualisiert zweidimensionale Daten:

```python
abb5 = plt.figure()
axes5 = abb5.add_subplot(1,1,1)

im5 = axes5.imshow(np.random.random((64, 64)), interpolation='none')
abb5.colorbar(im5)
```

```python
#axes5.imshow?
```

## Plots gestalten

Mit Matplotlib könnt ihr jegliche Plots erstellen, die ihr euch vorstellen könnt. Überlegt euch zuerst, wie der Plot aussehen soll. Die [Gallerie](http://matplotlib.org/gallery.html) kann dabei helfen. Mit einer beherzten Websuche, die meist direkt zur [Dokumentation](http://matplotlib.org/api/pyplot_api.html) führt, findet ihr dann die Funktionen, die ihr braucht.

> **Hinweis:** Plotten ist eine Kunst.

> **Weiterer Hinweis:** Versucht's mal mit `plt.xkcd()`.

### Titel, und Achsen und Legende mit LaTeX

Strings in Matplotlib wie Titel und Achsenbeschriftungen können **LaTeX-Code** enthalten. Text zwischen `$`-Zeichen wird dabei wie gewohnt im _Math-Mode_ gesetzt.

> **Achtung:** Einige Zeichen wie `\t` (`<TAB>`) oder `\n` (`<RETURN>`) werden von Python als Steuerzeichen und daher nicht als LaTeX-Code interpretiert. Markiert Strings, die LaTeX-Code enthalten, daher mit dem Prefix `r` (für _raw_), sodass Steuerzeichen ignoriert werden.

```python
abb6 = plt.figure()
axes6 = abb6.add_subplot(1,1,1)

x = np.linspace(0, 2 * np.pi, 100)
# Plot mit Label für Legende
axes6.plot(x, np.sin(x), label=r'$A \times \sin(\phi)$')
axes6.plot(x, np.cos(x), label=r'$A \times \cos(\phi)$')
# Titel
axes6.set_title('Oszillation')
# Achsenlimits
axes6.set_xlim(0, 2 * np.pi)
axes6.set_ylim(-1, 1)
# Achsenbeschriftungen
axes6.set_xlabel(r'Winkel $\phi \, [\mathrm{rad}]$')
axes6.set_ylabel(r'Auslenkung $d \, [\mathrm{cm}]$')
# Legende
axes6.legend(loc='lower left')
```

## Mehrere Plots in einer Abbildung

Eine Abbildung kann ganz einfach in mehrere `subplots` aufgeteilt werden. Hier zeigt sich jetzt die Bedeutung der drei einsen, die wir als Parameter dem `add_subplot` Befehl übergeben haben:

```python
# Eine Abbildung mit 2x1 Subplots erstellen
abb7 = plt.figure()
axes7_1 = abb7.add_subplot(2,1,1) # Ersten von 2x1 Subplots erstellen
axes7_2 = abb7.add_subplot(2,1,2) # Zweiten von 2x1 Subplots erstellen
# In beiden Subplots plotten
x = np.linspace(0, 2 * np.pi, 100)
axes7_1.plot(x, np.sin(x), label=r'$\sin(\phi)$')
axes7_2.plot(x, np.cos(x), label=r'$\cos(\phi)$')
# Die Abbildung konfigurieren
abb7.suptitle('Oszillation')
axes7_1.set_xlim(0, 2 * np.pi)
axes7_2.set_xlabel(r'Winkel $\phi \, [\mathrm{rad}]$')
axes7_1.legend(loc='lower left')
axes7_2.legend(loc='lower left')
```

## Plots speichern

Mit `matplotlib.figure.Figure.savefig` könnt ihr einen Plot als Bilddatei speichern:

```python
abb7.savefig('plots/my_plot.png')
```

> **Hinweis:** Um einen Plot im **DIN A4-Format** zu speichern könnt ihr dessen Größe und Auflösung anpassen:
>
> ```python
> abb7.set_size_inches(11.69, 8.27)
> abb7.savefig(filename, dpi=150)
> ```

> **Hinweis:** Um Plots in **LaTeX Dokumente** einzubinden bietet sich das [PGF Vektorformat](https://en.wikipedia.org/wiki/PGF/TikZ) statt Pixelgrafiken wie PNG an. Damit übernimmt LaTeX das Zeichnen des Plots und setzt zudem die enthaltenen Texte selbst, sodass die Fonts und Stile des Dokuments auch im Plot verwendet werden.
>
> Eine Beispiel-Implementierung findet ihr im [TexFig](https://github.com/knly/texfig) Repository.

## Aufgabe 1 - Temperaturen in Heidelberg visualisieren

Wir visualisieren nun die Temperaturdaten mit denen wir schon numerisch gearbeitet haben. Lest die Daten zunächst erneut ein:

```python
# Zelle ausführen, um Daten einzulesen
data = np.loadtxt('data/temperatures.txt')
date, T = data[np.abs(data[:,1]) != 99,:].transpose()
```

a) Plottet den Temperaturverlauf über den gesamten Zeitraum. Vergesst nicht Titel, Achsenbeschriftungen und Legende.

**Hinweis:** Ihr könnt die Temperatureinheit als `^\circ{}\mathrm{C}` (rendert als $^\circ{}\mathrm{C}$) schreiben.

<!-- grade_id: 202-1a-sol, points: 0, grade: false, solution: true -->
```python
aufg1aAbb = plt.figure()
aufg1aAxes = aufg1aAbb.add_subplot(1,1,1)

aufg1aAxes.plot(date, T, label='Messwerte')
aufg1aAxes.set_title('Temperaturverlauf in Heidelberg')
aufg1aAxes.set_xlim(np.min(date), np.max(date))
aufg1aAxes.set_xlabel(r'Zeitpunkt')
aufg1aAxes.set_ylabel(r'Temperatur $T \, [^\circ{}\mathrm{C}]$')
aufg1aAxes.legend()
```

Sieht dein Plot etwa so aus?

![Temperaturverlauf in Heidelberg](plots/202-1a.png)

```python
# Setze die Variable auf `True` wenn du mit deinem Plot zufrieden bist:
plot_1a_fertig = False
```

<!-- grade_id: 202-1a-test, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_true
assert_true(plot_1a_fertig, "Versuche dich an Aufgabe 1a, bis du mit deinem Plot zufrieden bist.")
print("Sieht gut aus.")
```

b) Plottet den kombinierten Jahres-Temperaturverlauf, also die Temperaturdaten aufgetragen über dem Zeitpunkt innerhalb des Jahres.

Plottet nur die Datenpunkte und keine verbindende Linie. Vergesst nicht Titel, Achsenbeschriftungen und Legende.

**Hinweis:** Den Zeitpunkt innerhalb eines Jahres erhaltet ihr wieder mit dem Modulo Operator: `date % 1`

<!-- grade_id: 202-1b-sol, points: 0, grade: false, solution: true -->
```python
aufg1bAbb = plt.figure()
aufg1bAxes = aufg1bAbb.add_subplot(1,1,1)

aufg1bAxes.scatter(date % 1, T, marker='.', label='Messwerte')
aufg1bAxes.set_title('Jahres-Temperaturverlauf in Heidelberg')
aufg1bAxes.set_xlim(0, 1)
aufg1bAxes.set_xlabel(r'Zeitpunkt innerhalb des Jahres')
aufg1bAxes.set_ylabel(r'Temperatur $T \, [^\circ{}\mathrm{C}]$')
aufg1bAxes.legend()
```

Sieht dein Plot etwa so aus?

![Jahres-Temperaturverlauf in Heidelberg](plots/202-1b.png)

```python
# Setze die Variable auf `True` wenn du mit deinem Plot zufrieden bist:
plot_1b_fertig = False
```

<!-- grade_id: 202-1b-test, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_true
assert_true(plot_1b_fertig, "Versuche dich an Aufgabe 1b, bis du mit deinem Plot zufrieden bist.")
print("🖼 Wunderschön.")
```

---

Nun kannst du Daten mit Numpy analysieren und mit Matplotlib plotten. Lerne in der nächsten Lektion wie du mit dem umfassenden wissenschaftlichen Paket _Scipy_ unter anderem Daten an Funktionen fitten kannst.

[Startseite](index.ipynb) | [**>> 203 - Fits mit Scipy**](203 - Fits mit Scipy.ipynb)